
<div class="alert alert-block alert-info">
   <h2 align="center"> <center><b> Reinforcement Learning Assignment 4 - Function Approximators Part 2: Deep Q Network </b></center></h2>

<br>
<center><font size="3">This notebook is a part of teaching material for ELEC-E8125</font></center>
<center><font size="3">Aalto University</font></center>
</div>


<a id='TOC'></a>

<div class="alert alert-block alert-success">
    
# Table of contents
* <a href='#1.'> 1. Introduction </a>
* <a href='#1.1'> 1.1 Task environments </a>
* <a href='#1.2'> 1.2 Learning Objectives </a>
* <a href='#1.3'> 1.3 Code Structure & Files </a>
* <a href='#1.4'> 1.4 Execution time </a>
* <a href='#2.'> 2. A (not-so-)deep Q-network</a>
* <a href='#3.'> 3. Submitting </a>
* <a href='#3.1'> 3.1 Feedback </a>

<a href='#T1'><b>Student Task 1.</b> Implementing DQN (10 points) </a>\
<a href='#Q1'><b>Student Question 1.1</b> Considering Continuous Action Spaces (5 points) </a>\
<a href='#Q2'><b>Student Question 1.2</b> Continuous Action Spaces Part 1 (15 points) </a>\
<a href='#Q3'><b>Student Question 1.3</b> The Target Network (10 points) </a>
 

**Total Points:** 40

**Estimated runtime of all the cells:** 2 hours

# 1. Introduction <a id='1.'></a>

In various real-world scenarios, dealing with high-dimensional state spaces makes it impractical to compute and store Q-values for every possible state-action pair in a Q-table. To address this challenge, we turn to function approximators. In this assignment, you will progress creating a basic Deep Q-Network (DQN) in the **Cartpole** and **LunarLander** environment.

**Please start working on this assignment early since the DQN will take some time to train.**


<div style="display:flex">
     <div style="flex:1;padding-left:100px;">
          <img src="imgs/cartpole.png" width="300"/>
         <figcaption style="flex:1;padding-left:20px;">  Figure 1: The Cartpole environment. </figcaption>
     </div>
     <div style="flex:1;padding-left:70px;">
          <img src="imgs/lunar_lander.png" width="300"/>
         <figcaption style="flex:1;padding-left:20px;">  Figure 2: The Lunarlander environment. </figcaption>
     </div>

</div>

Useful Sources:

- Mnih, Volodymyr, et al. "Playing atari with deep reinforcement learning." arXiv preprint arXiv:1312.5602 (2013). https://arxiv.org/pdf/1312.5602.pdf


## 1.1 Task environments: <a id='1.1'></a>
In this excercise, we will mainly use DQN for two tasks:
- Cartpole(https://gymnasium.farama.org/environments/classic_control/cart_pole/): This environment corresponds to the version of the cart-pole problem described by Barto, Sutton, and Anderson in “Neuronlike Adaptive Elements That Can Solve Difficult Learning Control Problem”. A pole is attached by an un-actuated joint to a cart, which moves along a frictionless track. The pendulum is placed upright on the cart and the goal is to balance the pole by applying forces in the left and right direction on the cart.

- Lunar Lander (https://gymnasium.farama.org/environments/box2d/lunar_lander/): This environment is a classic rocket trajectory optimization problem. According to Pontryagin’s maximum principle, it is optimal to fire the engine at full throttle or turn it off. This is the reason why this environment has discrete actions: engine on or off.

## 1.2 Learning Objectives: <a id='1.1'></a>
- Understand why and how function approximators can be used for Q-learning
- Understand the Deep Q-Network RL algorithm

## 1.3 Code Structure & Files <a id='1.2'></a>

```ex4_dqn.ipynb``` is the file needed to be modified for this part of the assignment.  

<span style="color:red"> **# IMPORTANT: DO NOT FORGET ANOTHER TASK IN ```ex4_rbf.ipynb```** </span>

```
├───cfg                            # Config files for environments
├───imgs                           # Images used in notebook
├───results
│   └───CartPole-v1
│   │   ├───logging                
│   │   │    └───logging.pkl        # Contains logged data
│   │   ├───model              
│   │   │    └───*dqn.pt            # Contains trained model
│   │   └───video                   # Videos saved
│   │   └───*cartpole_dqn.png       # Contains training performance plot
│   └───LunarLander-v2
│   │   ├───logging                
│   │   │    └───logging.pkl        # Contains logged data
│   │   ├───model              
│   │   │    └───*dqn.pt            # Contains trained model
│   │   └───video                   # Videos saved
│   │   └───*lunarlander_dqn.png    # Contains training performance plot
│   ex4_dqn.ipynb                   # 2nd assignment file containing tasks <---------This task
│   ex4_rbf.ipynb                   # 1st assignment file containing tasks <---------
│   train.py                        # Contains train and test functions 
│   utils.py                        # Contains useful functions 
└───buffer.py                       # Contains buffer functions
```

## 1.4 Execution time <a id='1.4'></a>

The training of these methods might take more than 40 mins depends on the server. If you have problem of experiment running takes too much time, you can download the jupyter notebook and test it locally.

## Warnings:

- Don’t copy and paste cells within a notebook. This will mess up the tracking metadata and prevent autograding from working.
- Only add new cells using the '+' button in the upper toolbar and do not split cells.
- Be cautious about things such as copying the whole notebook to Colab to work on it. This has sometimes resulted in removing all notebook metadata, making autograding impossible.

# 2. A (not-so-)deep Q-network <a id='2.'></a>

<a id='T1'></a>
<div class=" alert alert-warning">
    <h3><b>Student Task 1.</b> Implementing DQN (10 points) </h3> 

Finish the incomplete code in DQNAgent (functions ```update``` and ```get_action```, marked with ```TODO```) to implement a DQN agent.

**See Figure 3 for an example training performance plot for cartpole. Save the training performance plots, and check if they are in the right place (the paths please refer to <a href='#3.'>Submitting<a>).**

<figure style="text-align: center">
<img src="imgs/dqn.png" width="300px">
<figcaption style="text-align: center"> Figure 3: The training performance plot for cartpole-dqn might look something like this. </figcaption>
</figure>
    
🔝	<a href='#TOC'><b>Table of Contents</b></a>
</div>


```python
skip_training = False  # Set this flag to True before validation and submission
```

<!-- grade_id: cell-6421cfde6887e896, points: 0, grade: true, solution: false -->
```python

```

```python
import sys, os
sys.path.insert(0, os.path.abspath(".."))
from typing import List
import numpy as np
import torch, random, copy, yaml, time
import torch.nn as nn
import torch.optim as optim
import torch.nn.functional as F
from pathlib import Path

import utils as u
import train as t

from IPython.display import Video # to display videos

device = 'cpu'
```

```python
def mlp(in_dim, mlp_dims: List[int], out_dim, act_fn=nn.ReLU, out_act=nn.Identity):
    """Returns an MLP."""
    if isinstance(mlp_dims, int): raise ValueError("mlp dimensions should be list, but got int.")

    layers = [nn.Linear(in_dim, mlp_dims[0]), act_fn()]
    for i in range(len(mlp_dims)-1):
        layers += [nn.Linear(mlp_dims[i], mlp_dims[i+1]), act_fn()]
    # the output layer
    layers += [nn.Linear(mlp_dims[-1], out_dim), out_act()]
    return nn.Sequential(*layers)
```

```python
class DQNAgent(object):
    def __init__(self, state_shape, n_actions,
                 batch_size=32, hidden_dims=[12], gamma=0.98, lr=1e-3, grad_clip_norm=1000, tau=0.001):
        self.n_actions = n_actions
        self.state_dim = state_shape[0]
        self.policy_net = mlp(self.state_dim, hidden_dims, n_actions).to(device)
        self.target_net = copy.deepcopy(self.policy_net)
        self.target_net.eval()
        self.optimizer = optim.Adam(self.policy_net.parameters(), lr=float(lr))

        self.batch_size = batch_size
        self.gamma = gamma
        self.grad_clip_norm = grad_clip_norm
        self.tau = tau
        
        self.counter = 0

    def update(self, buffer):
        """ One gradient step, update the policy net."""
        start = time.perf_counter()
        self.counter += 1
        # Do one step gradient update
        batch = buffer.sample(self.batch_size, device=device)
        # Hint:
        #    state = batch.state
        #    action = batch.action 
        #    next_state = batch.next_state
        #    reward = batch.reward 
        #    not_done = batch.not_done 
        
        # TODO: Task 3: Finish the DQN implementation.
        # Hints: 1. You can use torch.gather() to gather values along an axis specified by dim. 
        #        2. torch.max returns a namedtuple (values, indices) where values is the maximum 
        #           value of each row of the input tensor in the given dimension dim.
        #           And indices is the index location of each maximum value found (argmax).
        #        3.  batch is a namedtuple, which has state, action, next_state, not_done, reward
        #           you can access the value be batch.<name>, e.g, batch.state
        #        4. check torch.nn.utils.clip_grad_norm_() to know how to clip grad norm
        #        5. You can go throught the PyTorch Tutorial given on MyCourses if you are not familiar with it. 
        # calculate the q(s,a)
        ########## You code starts here #########
     
        
        ########## You code ends here #########

        # update the target network
        u.soft_update_params(self.policy_net, self.target_net, self.tau)
        
        end = time.perf_counter()
        update_time = end - start
        return {'loss': loss.item(), 
                'q_mean': qs.mean().item(),
                'num_update': self.counter,
                'update_time': update_time}


    @torch.no_grad()
    def get_action(self, state, epsilon=0.05):
        # TODO:  Task 3: implement epsilon-greedy action selection
        ########## You code starts here #########
        p = np.random.uniform(0, 1)
        if p < epsilon:
            action = ...
        else:
            action = ...       
        ########## You code ends here #########
        return action


    def save(self, fp):
        path = fp/'dqn.pt'
        torch.save({
            'policy': self.policy_net.state_dict(),
            'policy_target': self.target_net.state_dict()
        }, path)

    def load(self, fp):
        path = fp/'dqn.pt'
        d = torch.load(path)
        self.policy_net.load_state_dict(d['policy'])
        self.target_net.load_state_dict(d['policy_target'])
```

```python
# init agent
with open(Path().cwd()/'cfg'/'cartpole_dqn.yaml', 'r') as f:
    cfg = t.Struct(**yaml.safe_load(f))
    
agent = DQNAgent(state_shape=cfg.state_shape, n_actions=cfg.n_actions, batch_size=cfg.batch_size, hidden_dims=cfg.hidden_dims,
                 gamma=cfg.gamma, lr=cfg.lr, tau=cfg.tau)
```

```python
# Trainig takes approximately 40 mins,depends on server
if not skip_training:
    t.train(agent,  cfg_path=Path().cwd()/'cfg'/'cartpole_dqn.yaml',) 
```

```python
# Visualize the DQN training plots for the cartpole 
if not skip_training:
    t.plot(save_name='cartpole_dqn.png', cfg_path=Path().cwd()/'cfg'/'cartpole_dqn.yaml',)
```

In order to get full points from the task make sure most of the test episodes give a reward of 200, given the seeds below (at least 50% of episodes).

```python
if not skip_training:
    t.test(agent, cfg_path=Path().cwd()/'cfg'/'cartpole_dqn.yaml', cfg_args=dict(save_video=True), seeds=[67, 23, 89, 12, 45, 78, 34, 90, 21, 56])
```

<!-- grade_id: cell-3ef5a6f870ccd026, points: 5, grade: true, solution: false -->
```python
%%capture --no-stdout
"TEST"
```

```python
if not skip_training:
      video = Video(Path().cwd()/'results'/'CartPole-v1'/'video'/'test'/'ex4_dqn-episode-0.mp4',
      embed=True, html_attributes="loop autoplay") # Set html_attributes="controls" for video control
      display(video)
```

```python
# init agent
with open(Path().cwd()/'cfg'/'lunarlander_dqn.yaml', 'r') as f:
    cfg = t.Struct(**yaml.safe_load(f))
    
agent = DQNAgent(state_shape=cfg.state_shape, n_actions=cfg.n_actions, batch_size=cfg.batch_size, hidden_dims=cfg.hidden_dims,
                 gamma=cfg.gamma, lr=cfg.lr, tau=cfg.tau)
```

```python
# Trainig takes approximately 60 mins,depends on server
if not skip_training:
    t.train(agent, cfg_path=Path().cwd()/'cfg'/'lunarlander_dqn.yaml', cfg_args=dict()) 
```

```python
# Visualize the DQN training plots for the lunarlander task
if not skip_training:
    t.plot(save_name='lunarlander_dqn.png', cfg_path=Path().cwd()/'cfg'/'lunarlander_dqn.yaml', cfg_args=dict())
```

In order to get full points from the task make sure most of the test episodes give a reward of 200, given the seeds below (at least 50% of episodes).

```python
if not skip_training:
    t.test(agent, cfg_path=Path().cwd()/'cfg'/'lunarlander_dqn.yaml', cfg_args=dict(save_video=True),seeds=[67, 23, 89, 12, 45, 78, 34, 90, 21, 56])
```

<!-- grade_id: cell-7c41d50955488cf2, points: 5, grade: true, solution: false -->
```python
%%capture --no-stdout
"TEST"
```

```python
if not skip_training:
  video = Video(Path().cwd()/'results'/'LunarLander-v2'/'video'/'test'/'ex4_dqn-episode-8.mp4',
  embed=True, html_attributes="loop autoplay") # Set html_attributes="controls" for video control
  display(video)
```

<a id='Q1'></a>
<div class=" alert alert-warning">
    <h3><b>Student Question 1.1</b> Considering Continuous Action Spaces (5 points) </h3> 
    
🔝	<a href='#TOC'><b>Table of Contents</b></a>
</div>


### 1.1.1 Question (5 points):
Can Q-learning be used directly in environments with continuous action spaces?

**Select one:**

1. Yes, Q-learning can be used directly without any modifications.
2. No, Q-learning cannot be used directly because it only works with a discrete set of actions.
3. No, Q-learning cannot be used directly in continuous action spaces, but it can be adapted with certain modifications.
4. Yes, Q-learning can be used directly, but it is less efficient than in discrete action spaces.

```python
sq_1_1_1 = 0 # Replace with the appropriate answer number
```

The following cells are used for grading

<!-- grade_id: cell-9b579c6ba03d991a, points: 0, grade: true, solution: false -->
```python
assert sq_1_1_1 in range(1, 5)
```

<!-- grade_id: cell-254ea90497e5d82c, points: 5, grade: true, solution: false -->
```python

```

<a id='Q2'></a>
<div class=" alert alert-warning">
    <h3><b>Student Question 1.2</b> Continuous Action Spaces Part 1 (15 points) </h3> 
    
🔝	<a href='#TOC'><b>Table of Contents</b></a>
</div>


### 1.2.1 Question (5 points):

What is the primary challenge when applying DQN to a continuous action space? Hint: select the one that describes the main issue.

**Select one:**

1. Updating the Q values
2. Storing the Q values
3. Performing the maximization (argmax) operation
4. Calculating the loss function
5. Calculating the reward

```python
sq_1_2_1 = 0 # Replace with the appropriate answer number
```

### 1.2.2 Question (5 points):

Which one of the following methods does NOT address the continuous action space problem in DQN?

**Select one:**

1. Learn a function approximation for the maximum of the Q-function
2. Use a policy to draw samples of Q(s,a) and select the action with the highest Q-value (e.g., MPO)
3. Approximate the Q-function as a convex function
4. Use a Deep Deterministic Policy Gradient (DDPG) approach to learn a policy and Q-function simultaneously
5. Use a Double DQN approach with two separate networks to estimate Q-values for continuous actions
6. Discretize the action space with fine-grained intervals

```python
sq_1_2_2 = 0 # Replace with the appropriate answer number
```

### 1.2.3 Question (5 points):

Discretization is one approach to address the continuous action space problem. What is a **significant** limitation of discretizing the action space to address the continuous action space problem in DQN?

**Select one:**

1. It may not generalize well to actions outside the discrete set
2. It increases the size of the state space exponentially
3. It requires modifying the neural network architecture
4. It can lead to unstable learning in the early stages of training
5. It makes the reward function non-differentiable

```python
sq_1_2_3 = 0 # Replace with the appropriate answer number
```

The following cells are used for grading

<!-- grade_id: cell-b7c78664ec74ec2a, points: 0, grade: true, solution: false -->
```python
assert sq_1_2_1 in range(1, 6)
assert sq_1_2_2 in range(1, 7)
assert sq_1_2_3 in range(1, 6)
```

<!-- grade_id: cell-cee299bf0c9ac8e4, points: 5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-bd84024a26ba389c, points: 5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-42e6f022fdf27684, points: 5, grade: true, solution: false -->
```python

```

<a id='Q3'></a>
<div class=" alert alert-warning">
    <h3><b>Student Question 1.3</b> The Target Network (10 points) </h3> 
    
🔝	<a href='#TOC'><b>Table of Contents</b></a>
</div>


### 1.3.1 Question (5 points):

Why do we use a separate target network in DQN instead of using the same network for both Q(s,a) and max_a(Q(s', ·))?

**Select one:**

1. To prevent the network from overfitting to the most recent transitions
2. To ensure that the Q-values are always within a specific range
3. To stabilize the training process by preventing frequent changes in target values
4. To enable the network to learn from both on-policy and off-policy data

```python
sq_1_3_1 = 0 # Replace with the appropriate answer number
```

### 1.3.2 Question (5 points):

What would happen if we don't stop the gradient of the target Q-value in DQN?

**Select one:**

1. The optimization process would become more stable due to the joint updates
2. The target network updates would be synchronized with the main network updates
3. Both the current Q-values and the target Q-values would be updated during optimization
4. The algorithm would effectively switch to an on-policy learning method

```python
sq_1_3_2 = 0 # Replace with the appropriate answer number
```

The following cells are used for grading

<!-- grade_id: cell-7427871eee3676e5, points: 0, grade: true, solution: false -->
```python
assert sq_1_3_1 in range(1, 5)
assert sq_1_3_2 in range(1, 5)
```

<!-- grade_id: cell-06da3985fcfd8e51, points: 5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-a8c5938cdf0dfd2c, points: 5, grade: true, solution: false -->
```python

```

# 3. Submitting <a id='3.'></a>
Ensure all tasks and questions (in ```ex4_dpn.ipynb``` and ```ex4_rbf.ipynb```) are answered and and that the necessary plots are saved in the appropriate locations. The relevant plots and files needed to be submitted for this assignment are:


- Training performance plots:
  - `cartpole_dqn.png`: Cartpole, training performance plots in terms of episode and episodic reward
  - `lunarlander_dqn.png`: Lunarlander, training performance plots in terms of episode and episodic reward 
<br>
<br>

  

- Model files:
  - `dqn.pt`: Trained model


Ensure the model files and plots are saved in correct paths:
- ```results/CartPole-v1/cartpole_dqn``` Training result for Cartpole environment
- ```results/CartPole-v1/model/dqn.pt``` Model for Cartpole environment
- ```results/LunarLander-v2/lunarlander_dqn.png``` Training result for LunarLander environment
- ```results/LunarLander-v2/model/dqn.pt``` Model for LunarLander environment


<span style="color:red"> **# IMPORTANT: DO NOT FORGET ANOTHER TASK IN ```ex4_rbf.ipynb```** </span>

<!-- grade_id: cell-c97f1e23824070fc, points: 0, grade: true, solution: false -->
```python
# Make sure that skip training is set to True before submission
assert skip_training == True
```

## 3.1 Feedback <a id='3.1'></a>

In order to help the staff of the course as well as the forthcoming students, it would be great if you could answer to the following questions in your submission:

1) How much time did you spend solving this exercise? (change the ```hrs``` variable below to a floating point number representing the number of hours taken e.g. 5.43)

```python
hrs = None
```

2) Difficulty of each task/question from 1-5 (int or float)

```python
T1 = None   # Implementing DQN (10 points)
Q1_1 = None # Question 1.1 Considering Continuous Action Spaces (5 points)
Q1_2 = None # Question 1.2 Continuous Action Spaces Part 1 (15 points)
Q1_3 = None # Question 1.3 The Target Network (10 points)
```

3) How well did you understand the content of the task/question from 1-5? (int or float)

```python
T1 = None   # Implementing DQN (10 points)
Q1_1 = None # Question 1.1 Considering Continuous Action Spaces (5 points)
Q1_2 = None # Question 1.2 Continuous Action Spaces Part 1 (15 points)
Q1_3 = None # Question 1.3 The Target Network (10 points)
```

4) General feedback. Consider questions like:

    - Did the content of the lecture relate well with the assignment?
    - To what extent did you find the material to be potentially useful for your research and studies?
    
Please share any additional feedback, suggestions, or comments you have about the lecture, assignment, or course content. Your input is valuable in helping us improve the learning experience.

DOUBLE CLICK HERE TO EDIT, CLEAR THIS TEXT AND ANSWER HERE